rule_files:
  - prometheus.yaml   # The file containing your Velero alert rule

tests:
  # ───────────────────────────────────────────────
  # (1) Velero pod is UP but no successful backups.
  #     -> Alert should fire.
  # ───────────────────────────────────────────────
  - interval: 5m
    input_series:
      - series: 'up{pod="velero"}'
        values: '1+0x4'  # Velero pod active
      - series: 'velero_backup_success_total{schedule="daily"}'
        values: '0+0x4'  # No successful backups yet
    alert_rule_test:
      - eval_time: 15m
        alertname: VeleroNoFirstSuccessfulBackup
        exp_alerts:
          - exp_labels:
              severity: critical
              schedule: daily
            exp_annotations:
              summary: "Velero backup has not completed once successfully for any schedule."
              description: "No successful Velero backups have been detected for schedule: daily while the Velero pod is running.This likely indicates a configuration or permission issue preventing backups from completing."

  # ───────────────────────────────────────────────
  # (2) Velero pod is DOWN — no alert should fire.
  # ───────────────────────────────────────────────
  - interval: 5m
    input_series:
      - series: 'up{pod="velero"}'
        values: '0+0x4'  # Velero pod down
      - series: 'velero_backup_success_total{schedule="daily"}'
        values: '0+0x4'
    alert_rule_test:
      - eval_time: 15m
        alertname: VeleroNoFirstSuccessfulBackup
        exp_alerts: []   # No alerts expected when pod is down

  # ───────────────────────────────────────────────
  # (3) Velero pod UP and first successful backup done.
  #     -> Alert should clear.
  # ───────────────────────────────────────────────
  - interval: 5m
    input_series:
      - series: 'up{pod="velero"}'
        values: '1+0x4'
      - series: 'velero_backup_success_total{schedule="daily"}'
        values: '0+0x2 1+0x2'  # Backup success after 10 minutes
    alert_rule_test:
      - eval_time: 20m
        alertname: VeleroNoFirstSuccessfulBackup
        exp_alerts: []   # Alert clears once success recorded

  # ───────────────────────────────────────────────
  # (4) Velero pod UP but backup metric missing (metric drop case or metric not initialized case)
  #     -> Alert should fire since no success metric found.
  # ───────────────────────────────────────────────
  - interval: 5m
    input_series:
      - series: 'up{pod="velero"}'
        values: '1+0x4'
    alert_rule_test:
      - eval_time: 15m
        alertname: VeleroNoFirstSuccessfulBackup
        exp_alerts:
          - exp_labels:
              severity: critical
            exp_annotations:
              summary: "Velero backup has not completed once successfully for any schedule."
              description: "Velero backup metrics are missing while the Velero pod is running. This may indicate that the backup metric is not being emitted or there is a configuration/permission issue."
